---
title       : Basic operations
description : >-
  This chapter shows you how to use Git to back up your work, and how
  to view the history of what you have done.

--- type:PureMultipleChoiceExercise lang:bash xp:50 key:a3732cc273
## What is version control?

A **version control system** is a tool that manages changes made to the files and directories in a project.
Many version control systems exist;
this lesson focuses on one called Git,
which is used by many of the data science tools covered in our other lessons.
Its strengths are:

* Nothing that is committed to Git is ever lost,
  so you can always go back to see which results were generated by which versions of your programs.

* Git automatically notifies you when your work conflicts with someone else's,
  so it's harder (but not impossible) to accidentally overwrite work.

* Git can synchronize work done by different people on different machines,
  so it scales as your team does.

Version control isn't just for software:
books, papers, data sets, and anything that changes over time or needs to be shared
can and should be stored and shared using something like Git.

<hr>

Which of the following does Git do?

*** =possible_answers
- Keep track of changes to files.
- Notice conflicts between changes made by different people.
- Synchronize files between different computers.
- [All of the above.]

*** =hint

*** =feedbacks
- Yes, but that's not all.
- Yes, but that's not all.
- Yes, but that's not all.
- Correct.

<!-- -------------------------------------------------------------------------------- -->

--- type:PureMultipleChoiceExercise lang:bash xp:50 key:74e878cc1b
## How does Git compare to alternatives?

Git and other version control systems are the only way to track and share work.
Some popular alternatives are:

1. *Email.*
   Lots of people mail files back and forth,
   but it's hard to keep track of which version is the most recent,
   and even harder to be sure that everyone's edits have been merged together.

2. *Putting files in a shared network drive or on Dropbox.*
   This is more convenient than email,
   but still makes people do the work of tracking changes.

3. *Online editing.*
   Tools like Google Docs work well for small groups,
   but it's hard to keep track of changes as the number of contributors and comments grows.

<hr>

What kind of work should you store in version control?

*** =possible_answers
- Software.
- Small datasets.
- Reports and other written work.
- [All of the above.]

*** =hint

*** =feedbacks
- Yes, but that's not all.
- Yes, but that's not all.
- Yes, but that's not all.
- Correct!

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:1dac2b0a28
## Where does Git store information?

Each of your Git projects has two parts:
the files and directories that you create and edit directly,
and the extra information that Git records about the project's history.
The combination of these two things is called a **repository**.

Git stores all of its extra information in a directory called `.git`
located in the root directory of the repository.
Git expects this information to be laid out in a very precise way,
so you should never edit or delete anything in `.git`.

<hr>

Suppose your home directory `/home/repl` contains a repository called `dental`,
which has a sub-directory called `data`.
Where is information about the history of the files in `/home/repl/dental/data` stored?

*** =instructions
- `/home/repl/.git`
- [`/home/repl/dental/.git`]
- `/home/repl/dental/data/.git`
- None of the above.

*** =hint

Use `ls -a` to see all of the contents of a directory.

*** =pre_exercise_code
```{shell}

```

*** =sct
```{python}
e1 = 'No: every repository stores its own history.'
c2 = 'Yes: all of the information about a repository is stored under its root directory.'
e3 = 'No: all of the information about a repository is stored under its root directory.'
e4 = 'No: one of the answers above is correct.'
Ex() >> test_mc(2, [e1, c2, e3, e4])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:97ae2f8613
## How can I view a repository's entire history?

Every Git command consists of the word `git` followed by a verb that specifies what you want to do.
One of the most common commands is `git log`,
which tells Git to display the **log** of the project's history.
There are two ways to run it:

1. If you are in a project's root directory or one of its sub-directories,
   you can just run `git log`.
2. Otherwise, you can run `git -C /path/to/repository log`.
   The flag `-C /some/path` works with every Git command,
   which is why it comes between the word `git` and the specific command.

Log entries are shown most recent first,
and look like this:

```
commit 0430705487381195993bac9c21512ccfb511056d
Author: Rep Loop <repl@datacamp.com>
Date:   Wed Sep 20 13:42:26 2017 +0000

    Added year to report title.
```

Ignoring the `commit` line for a moment,
the other lines show who made the change,
when,
and what comment (called a **log message**) they wrote for the change.

<hr>

You are in the directory `dental`,
which is a Git repository.
Use a single Git command *without* the `-C` option
to view the repository's history.
What is the message on the very first entry in the log
(which is displayed last)?

*** =instructions
- "Added summary report file."
- "Adding seasonal CSV data files."
- "Fixed bug and generated results."
- "Adding reminder to cite funding sources."

*** =hint

Use spacebar to page down through the log and 'q' to quit.

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
err = 'Incorrect: please re-run the command and use `git log`.'
Ex() >> test_student_typed(r'\s*git\s+log\s*', \
                           fixed=False, \
                           msg='Use `git log`.') \
     >> test_mc(1, ['Correct', err, err, err])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:77aa525e25
## How can I view a specific file's history?

A project's entire log can be overwhelming,
so it's often useful to inspect only the changes to particular files or directories.
You can do this using `git log path`,
where `path` is the path to a specific file or directory.
The log for a file shows changes made to that file;
the log for a directory shows when files were added or deleted in that directory,
rather than when the contents of the directory's files were changed.

<hr>

You are in the `dental` repository.
Use `git log` to display only the changes made to `data/southern.csv`.
How many have there been?

*** =instructions
- 0.
- 1.
- 2.
- 3.

*** =hint

Use `git log` and the path to the file you are interested in.

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
err = 'Incorrect: please use `git log data/southern.csv` and count the number of log entries.'
Ex() >> test_student_typed(r'\s*git\s+log\s+data/southern.csv\s*', \
                           fixed=False, \
                           msg='Use `git log` and the name of a file.') \
     >> test_mc(3, [err, err, 'Correct!', err])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:d03c5fa3e6
## How can I view a specific log entry?

Every commit to a repository has a unique hexadecimal ID called a **hash**
(since it is generated by running the changes through a pseudo-random number generator
called a **hash function**).
To view the details of a specific commit,
you can use the command `git show` with the first few characters of the commit's hash.
(5 or 6 characters is usually enough to tell Git which hash you mean.)
For example, the command `git show 043070` produces this:

    commit 0430705487381195993bac9c21512ccfb511056d
    Author: Rep Loop <repl@datacamp.com>
    Date:   Wed Sep 20 13:42:26 2017 +0000
    
        Added year to report title.
    
    diff --git a/report.txt b/report.txt
    index e713b17..4c0742a 100644
    --- a/report.txt
    +++ b/report.txt
    @@ -1,4 +1,4 @@
    -# Seasonal Dental Surgeries 2017-18
    +# Seasonal Dental Surgeries (2017) 2017-18
     
     TODO: write executive summary.

The first part is the log entry you have seen before.
The second part shows what changes the commit made:
sections that the change removed are prefixed with `-`,
while sections that it added are prefixed with `+`.

<hr>

How many files were changed by the commit whose hash starts with `aa5991c6`?
(Remember, you only need to type the first few characters of the hash.)

*** =instructions
- None.
- 1.
- 2.
- 4.

*** =hint

Use `git show` and the first 6 characters of the hash
and then look at the names of the files that were changed.

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
e_more = 'No, there have been more changes than that.'
correct = 'Correct!'
e_fewer = 'No, there have been fewer changes than that.'
Ex() >> test_student_typed(r'\s*git\s+show\s+4d3e[0-9a-f]*\s*', \
                           fixed=False, \
                           msg='Use `git show` and the first few digits of the commit hash.') \
     >> test_mc(3, [e_more, e_more, correct, e_fewer])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:eef645517f
## What is in a diff?

A **diff** is a formatted display of the differences between two sets of files.
By default,
Git displays diffs like this:

    diff --git a/report.txt b/report.txt
    index e713b17..4c0742a 100644
    --- a/report.txt
    +++ b/report.txt
    @@ -1,4 +1,4 @@
    -# Seasonal Dental Surgeries 2017-18
    +# Seasonal Dental Surgeries (2017) 2017-18
     
     TODO: write executive summary.

This shows:
- The command `diff --git` used to produce the output.
  In it, `a` and `b` are placeholders meaning "the first version" and "the second version".
- An index line showing keys into Git's internal database of changes.
- `--- a/report.txt` and `+++ b/report.txt`,
  which indicate that lines being *removed* are prefixed with `-`,
  while lines being added are prefixed with `+`.
- A line starting with `@@` which tells where the changes are being made.
  Here, the line shows that lines 1-4 are being removed and replaced with new lines.
- A line-by-line listing of the changes.

Desktop programming tools like [RStudio](https://www.rstudio.com/) can turn diffs like this
into a more readable side-by-side display of changes;
you can also use standalone tools like [DiffMerge](https://sourcegear.com/diffmerge/)
or [WinMerge](http://winmerge.org/).

<hr>

You are in the `dental` repository.
Use `git diff` with a hash and a filename to look at the commit with ID `166a86`.
How many lines did it change in the file `bin/teeth`?

*** =instructions
- None.
- 1.
- 2.
- 20.

*** =hint

Use `git show`, the ID, and the filename.

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
err_some = 'No, the commit changed some of the lines in `bin/teeth`.'
correct = 'Yes, the commit changed one line.'
err_fewer = 'No, the commit did not change that many lines in `bin/teeth`.'
Ex() >> test_student_typed(r'\s*git\s+show\s+ed0ec0[0-9a-f]?\s*', \
                           fixed=False, \
                           msg='Use `git show ed0ec0` and examine the diff.') \
     >> test_mc(2, [err_some, correct, err_fewer, err_fewer])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:1bc30ab583
## What is Git's equivalent of a relative path?

A hash is like an absolute path:
it identifies a specific commit.
Another way to identify a commit is to use the equivalent of a relative path.
The special label `HEAD` always refers to the most recent commit;
the label `HEAD~1` then refers to the commit before it,
while `HEAD~2` refers to the commit before that,
and so on.
(Note that the symbol between `HEAD` and the number is a tilde `~`,
and *not* a minus sign `-`.)

<hr>

You are in the `dental` repository.
Using a single Git command,
look at the commit made just before the most recent one.
Which of the following files did it change?

*** =instructions
- `report.txt`.
- `data/western.csv`.
- Both of the above.
- Neither of the above.

*** =hint

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
err_and = 'Yes, but it also changed another file.'
correct = 'Correct.'
err_some = 'No, the commit `HEAD~1` did change some files.'
Ex() >> test_student_typed(r'\s*git\s+show\s+HEAD~1\s*', \
                           fixed=False, \
                           msg='Use `git show` and remember that `HEAD~N` is current minus that many.') \
     >> test_mc(3, [err_and, err_and, correct, err_some])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:f76c7f5965
## How can I see who changed what in a file?

`git log` displays the overall history of a project or file,
but Git can give even more information:
the command `git annotate file` shows who made the last change to each line of a file and when.
For example,
the first three lines of output from `git annotate report.txt` look something like this:

```
04307054        (  Rep Loop     2017-09-20 13:42:26 +0000       1)# Seasonal Dental Surgeries (2017) 2017-18
5e6f92b6        (  Rep Loop     2017-09-20 13:42:26 +0000       2)
5e6f92b6        (  Rep Loop     2017-09-20 13:42:26 +0000       3)TODO: write executive summary.
```

The first column is the hash of the most recent commit to change that line.
The other columns show who made the change,
the date and time it was made,
the line number,
and the line itself.

<hr>

You are in the `dental` repository.
Use a single command to see the changes to `report.txt`.
How many different sets of changes have been made to this file
(i.e., how many different hashes show up in the first column of the output)?

*** =instructions
- 1.
- 2.
- 3.
- 4.

*** =hint

Use `git annotate report.txt` and count the number of distinct commit IDs.

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
e_more = 'No, there have been more changes than that.'
correct = 'Correct!'
e_fewer = 'No, there have been fewer changes than that.'
Ex() >> test_student_typed(r'\s*git\s+annotate\s+report.txt\s*', \
                           fixed=False, \
                           msg='Use `git annotate` and the name of a file.') \
     >> test_mc(3, [e_more, e_more, correct, e_fewer])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:9862ae22bd
## How can I see what changed between two commits?

`git show` shows the changes made *in* a particular commit.
To see the changes *between* two widely-separated commits,
you can use `git diff R1..R2`,
where `R1` and `R2` identify the two commits you're interested in,
and the connector `..` is a pair of dots.
For example,
`git diff abc123..def456` shows the differences between the commits `abc123` and `def456`,
while `git diff HEAD~1..HEAD~3` shows the differences between the state of the repository one commit in the past
and its state three commits in the past.

<hr>

You are in the `dental` repository.
Use `git diff` to view the differences between its current state
and its state in the previous-but-two commit.
Which of the following files have changed?

*** =instructions
- `data/western.csv`.
- `report.txt`.
- `data/southern.csv`.
- `report.txt` and `data/western.csv`.
- `report.txt` and `data/southern.csv`.

*** =hint

Use `git log` with two `HEAD`-based commit identifiers and look for filenames.

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
err_more = 'Yes, but another file was changed as well.'
err_not = 'No, that file did not change.'
correct = 'Correct!'
err_half = 'No, one of those files did not change.'
Ex() >> test_student_typed(r'\s*git\s+diff\s-r\s+(HEAD\.\.HEAD~2)|(HEAD\.\.HEAD~2)\s*', \
                           fixed=False, \
                           msg='Use `git diff` and remember that `HEAD~N` is current minus that many.') \
     >> test_mc(4, [err_more, err_more, err_not, correct, err_half])
```
