---
title       : Basic workflow
description : >-
  This chapter explains what version control is and why you should use
  it, and introduces the most common steps in a common Git workflow.

--- type:PureMultipleChoiceExercise lang:bash xp:50 key:a3732cc273
## What is version control?

A **version control system** is a tool that manages changes made to the files and directories in a project.
Many version control systems exist;
this lesson focuses on one called Git,
which is used by many of the data science tools covered in our other lessons.
Its strengths are:

* Nothing that is committed to Git is ever lost,
  so you can always go back to see which results were generated by which versions of your programs.

* Git automatically notifies you when your work conflicts with someone else's,
  so it's harder (but not impossible) to accidentally overwrite work.

* Git can synchronize work done by different people on different machines,
  so it scales as your team does.

Version control isn't just for software:
books, papers, data sets, and anything that changes over time or needs to be shared
can and should be stored and shared using something like Git.

<hr>

Which of the following does Git do?

*** =possible_answers
- Keep track of changes to files.
- Notice conflicts between changes made by different people.
- Synchronize files between different computers.
- [All of the above.]

*** =hint

*** =feedbacks
- Yes, but that's not all.
- Yes, but that's not all.
- Yes, but that's not all.
- Correct.

<!-- -------------------------------------------------------------------------------- -->

--- type:PureMultipleChoiceExercise lang:bash xp:50 key:74e878cc1b
## How does Git compare to alternatives?

Git and other version control systems are the only way to track and share work.
Some popular alternatives are:

1. *Email.*
   Lots of people mail files back and forth,
   but it's hard to keep track of which version is the most recent,
   and even harder to be sure that everyone's edits have been merged together.

2. *Putting files in a shared network drive or on Dropbox.*
   This is more convenient than email,
   but still makes people do the work of tracking changes.

3. *Online editing.*
   Tools like Google Docs work well for small groups,
   but it's hard to keep track of changes as the number of contributors and comments grows.

<hr>

What kind of work should you store in version control?

*** =possible_answers
- Software.
- Small datasets.
- Reports and other written work.
- [All of the above.]

*** =hint

*** =feedbacks
- Yes, but that's not all.
- Yes, but that's not all.
- Yes, but that's not all.
- Correct!

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:1dac2b0a28
## Where does Git store information?

Each Git project has two parts:
the files and directories that you create and edit directly,
and the extra information that Git records about the project's history.
The combination of these two things is called a **repository**.

Git stores all of its extra information in a directory called `.git`
located in the root directory of the repository.
Git expects this information to be laid out in a very precise way,
so you should never edit or delete anything in `.git`.

<hr>

Your home directory `/home/repl` contains a repository called `dental`,
which has a sub-directory called `data`.
Where is information about the history of the files in `/home/repl/dental/data` stored?

*** =instructions
- `/home/repl/.git`
- [`/home/repl/dental/.git`]
- `/home/repl/dental/data/.git`
- None of the above.

*** =hint

Use `ls -a` to see all of the contents of a directory.

*** =pre_exercise_code
```{shell}

```

*** =sct
```{python}
e1 = 'No: every repository stores its own history.'
c2 = 'Yes: all of the information about a repository is stored under its root directory.'
e3 = 'No: all of the information about a repository is stored under its root directory.'
e4 = 'No: one of the answers above is correct.'
Ex() >> test_mc(2, [e1, c2, e3, e4])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:bc52cf1174
## How can I check the state of a repository?

The first step in using Git is to check the **status** of your repository's files.
To do this,
run the command `git status`,
which displays a list of the files that have been modified
since the last time changes were committed.

<hr>

You are in the `dental` repository.
Use a single Git command to discover which file(s) have been changed since the last commit.
Which file(s) are listed?

*** =instructions
- `data/summer.csv`.
- `report.txt`.
- Neither of the above.
- Both of the above.

*** =hint

*** =pre_exercise_code
```{python}
append = '''

Fehrenbach: "Dental Anatomy Coloring Book" (2e), 2013.
'''
with open('dental/report.txt', 'w') as writer:
    writer.write(append)
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
Ex() >> test_student_typed(r'\s*git\s+status\s*', \
                           fixed=False, \
                           msg='Use `git status`.') \
     >> test_mc(2, ['No, that file has not changed.', \
                    'Correct!', \
                    'No, one file has changed.', \
                    'No, only one file has changed.'])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:BulletConsoleExercise key:0dd628a298
## How can I tell what I have changed?

Git has a **staging area** in which it stores files with changes you want to save
that haven't been saved yet.
`git status` shows you which files are in the staging area,
and which files have changes that haven't yet been put there.
In order to compare the file as it currently is to what you last saved,
you can use `git diff filename`.
`git diff` without any filenames will show you all the changes in your repository,
while `git diff directory` will show you the changes to the files in some directory.

*** =pre_exercise_code
```{python}
with open('dental/data/northern.csv', 'w') as writer:
    writer.write('2017-11-01,bicuspid\n')
repl = connect('bash')
repl.run_command('cd dental')
```

*** =type1: ConsoleExercise
*** =key1: e9c198755a

*** =xp1: 10

*** =instructions1

You are in the `dental` repository.
Use one Git command to see what has been changed.

*** =hint1

*** =sample_code1
```{shell}
```

*** =solution1
```{shell}
git status
```

*** =sct1
```{python}
Ex() >> test_student_typed(r'\s*git\s+diff.*',
                           fixed=False,
                           msg='Use `git diff`.')
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:eef645517f
## What is in a diff?

A **diff** is a formatted display of the differences between two sets of files.
By default,
Git displays diffs like this:

    diff --git a/report.txt b/report.txt
    index e713b17..4c0742a 100644
    --- a/report.txt
    +++ b/report.txt
    @@ -1,4 +1,4 @@
    -# Seasonal Dental Surgeries 2017-18
    +# Seasonal Dental Surgeries (2017) 2017-18
     
     TODO: write executive summary.

This shows:
- The command used to produce the output (in this case, `diff --git`).
  In it, `a` and `b` are placeholders meaning "the first version" and "the second version".
- An index line showing keys into Git's internal database of changes.
  We will explore these in the next chapter.
- `--- a/report.txt` and `+++ b/report.txt`,
  which indicate that lines being *removed* are prefixed with `-`,
  while lines being added are prefixed with `+`.
- A line starting with `@@` that tells where the changes are being made.
  Here, the line shows that lines 1-4 are being removed and replaced with new lines.
- A line-by-line listing of the changes.

Desktop programming tools like [RStudio](https://www.rstudio.com/) can turn diffs like this
into a more readable side-by-side display of changes;
you can also use standalone tools like [DiffMerge](https://sourcegear.com/diffmerge/)
or [WinMerge](http://winmerge.org/).

<hr>

You are in the `dental` repository.
Use `git diff` look at the changes to `dental/data/northern.csv`.
How many lines did it change in the file `bin/teeth`?

*** =instructions
- None.
- 1.
- 2.
- 20.

*** =hint

Use `git diff` and the filename.

*** =pre_exercise_code
```{shell}
with open('dental/data/northern.csv', 'w') as writer:
    writer.write('2017-11-01,bicuspid\n')
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
err_some = 'No, the commit changed some of the lines.'
correct = 'Yes, the commit changed one line.'
err_fewer = 'No, the commit did not change that many lines.'
Ex() >> test_mc(2, [err_some, correct, err_fewer, err_fewer])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:BulletConsoleExercise key:55204eee61
## What's the first step in saving changes?

You commit changes to a Git repository in two steps:
you add one or more files to the staging area,
and then you **commit** everything in the staging area.
Putting files in the staging area is like putting things in a box,
while committing is like putting it in the mail:
you can add more things to the box or take things out as often as you want,
but once you put it in the mail,
you can't make further changes.

To add a file to the staging area,
using `git add filename`.
You can do this several times in a row as you make changes to the file.

*** =pre_exercise_code
```{python}
append = '''

Fehrenbach: "Dental Anatomy Coloring Book" (2e), 2013.
'''
with open('dental/report.txt', 'w') as writer:
    writer.write(append)
connect('bash').run_command('cd dental')
```

*** =type1: ConsoleExercise
*** =key1: a34bf017f0

*** =xp1: 10

*** =instructions1

You are in the `dental` repository.
Use `git add` to add the file `report.txt` to the staging area.

*** =hint1

*** =sample_code1
```{shell}
```

*** =solution1
git add report.txt
```{shell}
```

*** =sct1
```{python}
Ex() >> test_student_typed(r'\s*git\s+add\s+report.txt\s*',
                           fixed=False,
                           msg='Use `git add` and a filename.')
```

*** =type2: ConsoleExercise
*** =key2: 961661800c

*** =xp2: 10

*** =instructions2

Use another Git command to check the repository's status.

*** =hint2

*** =sample_code2
```{shell}
```

*** =solution2
```{shell}
git status
```

*** =sct2
```{python}
Ex() >> test_student_typed(r'\s*git\s+status\s*',
                           fixed=False,
                           msg='Use `git status`.')
```

<!-- -------------------------------------------------------------------------------- -->

--- type:BulletConsoleExercise key:f208f45d7d
## How can I tell what's going to be committed?

In order to compare the file as it currently is
to the changes in the staging area,
you can use `git diff -r HEAD filename`.
The `-r` flag means "compare to a particular revision",
and `HEAD` is a shortcut meaning "the most recent commit".
We will explore other uses of `-r` and `HEAD` in the next chapter.

*** =pre_exercise_code
```{python}
with open('dental/data/northern.csv', 'w') as writer:
    writer.write('2017-11-01,bicuspid\n')
with open('dental/data/eastern.csv', 'w') as writer:
    writer.write('2017-11-02,molar\n')
repl = connect('bash')
repl.run_command('cd dental')
repl.run_command('git add data/northern.csv')
```

*** =type1: ConsoleExercise
*** =key1: 77c975a5c8

*** =xp1: 10

*** =instructions1

You are in the `dental` repository.
Use `git diff` with appropriate arguments to see how files differ from
the last saved revision.

*** =hint1

*** =sample_code1
```{shell}
```

*** =solution1
```{shell}
git diff -r HEAD
```

*** =sct1
```{python}
Ex() >> test_student_typed(r'\s*git\s+diff\s+-r\s+HEAD\s*',
                           fixed=False,
                           msg='Use `git diff -r HEAD (with or without a directory name)`.')
```

*** =type2: ConsoleExercise
*** =key2: 5a866b3ef0

*** =xp2: 10

*** =instructions2

Use a single Git command to view the changes in the file that has been staged
(and *only* that file),

*** =hint2

*** =sample_code2
```{shell}
```

*** =solution2
```{shell}
git diff -r HEAD data/northern.csv | cat
```

*** =sct2
```{python}
Ex() >> test_student_typed(r'\s*git\s+diff\s+-r\s+HEAD\s+data/northern\.csv\s*',
                           fixed=False,
                           msg='Use `git diff -r HEAD` and a filename.')
```

*** =type3: ConsoleExercise
*** =key3: 538a9c35d4

*** =xp3: 10

*** =instructions3

Use one Git command to add the other changed file to the staging area.

*** =hint3

*** =sample_code3
```{shell}
```

*** =solution3
```{shell}
git add data/eastern.csv
```

*** =sct3
```{python}
Ex() >> test_student_typed(r'\s*git\s+add\s+data/eastern\.csv\s*',
                           fixed=False,
                           msg='Use `git add`.')
```

<!-- -------------------------------------------------------------------------------- -->

--- type:BulletConsoleExercise key:dbf5aa722c
## How do I commit changes?

To save the changes in the staging area,
you use the command `git commit`.
It always saves everything that is in the staging area as one unit:
as you will see later,
when you want to undo changes to a project,
you undo all of a commit or none of it.

When you commit changes,
Git requires you to enter a **log message**.
This serves the same purpose as a comment in a program:
it tells the next person to examine the repository why you made a change.

By default,
Git launches a text editor to let you write this message.
To keep things simple,
you can use `-m "some message in quotes"` on the command line
to enter a single-line message like this:

```
git commit -m "Program appears to have become self-aware."
```

*** =pre_exercise_code
```{python}
append = '''

Fehrenbach: "Dental Anatomy Coloring Book" (2e), 2013.
'''
with open('dental/report.txt', 'w') as writer:
    writer.write(append)
repl = connect('bash')
repl.run_command('cd dental')
repl.run_command('git add report.txt')
```

*** =type1: ConsoleExercise
*** =key1: 17e219ea22

*** =xp1: 10

*** =instructions1

You are in the `dental` repository.
Use one Git command to check the status of the repository.

*** =hint1

*** =sample_code1
```{shell}
```

*** =solution1
```{shell}
git status
```

*** =sct1
```{python}
Ex() >> test_student_typed(r'\s*git\s+status\s*',
                           fixed=False,
                           msg='Remember, you want to check the *status* of the repository.')
```

*** =type2: ConsoleExercise
*** =key2: a5ce3bebb8

*** =xp2: 10

*** =instructions2

Commit the changes in the staging area with the message "Adding a reference."

*** =hint2

*** =sample_code2
```{shell}
```

*** =solution2
```{shell}
git commit -m "Adding a reference."
```

*** =sct2
```{python}
Ex() >> test_student_typed(r'\s*git\s+commit\s+-m\s+"Adding\s+a\s+reference."\s*',
                           fixed=False,
                           msg='Use `git commit` with `-m "message"`.')
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:97ae2f8613
## How can I view a repository's history?

The command `git log` is used to view the **log** of the project's history.
Log entries are shown most recent first,
and look like this:

```
commit 0430705487381195993bac9c21512ccfb511056d
Author: Rep Loop <repl@datacamp.com>
Date:   Wed Sep 20 13:42:26 2017 +0000

    Added year to report title.
```

The `commit` line displays a unique ID for the commit called a **hash**;
we will explore these further in the next chapter.
The other lines tell you who made the change,
when,
and what log message they wrote for the change.

When you run `git log`,
Git automatically uses a pager to show one screen of output at a time.
Press the space bar to go down a page or the 'q' key to quit.

<hr>

You are in the directory `dental`,
which is a Git repository.
Use a single Git command (*without* the `-C` option)
to view the repository's history.
What is the message on the very first entry in the log
(which is displayed last)?

*** =instructions
- "Added summary report file."
- "Adding seasonal CSV data files."
- "Fixed bug and generated results."
- "Adding reminder to cite funding sources."

*** =hint

Use the space bar to page down through the log and 'q' to quit.

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
err = 'Incorrect: please re-run the command and use `git log`.'
Ex() >> test_mc(1, ['Correct', err, err, err])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:77aa525e25
## How can I view a specific file's history?

A project's entire log can be overwhelming,
so it's often useful to inspect only the changes to particular files or directories.
You can do this using `git log path`,
where `path` is the path to a specific file or directory.
The log for a file shows changes made to that file;
the log for a directory shows when files were added or deleted in that directory,
rather than when the contents of the directory's files were changed.

<hr>

You are in the `dental` repository.
Use `git log` to display only the changes made to `data/southern.csv`.
How many have there been?

*** =instructions
- 0.
- 1.
- 2.
- 3.

*** =hint

Use `git log` and the path to the file you are interested in.

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
err = 'Incorrect: please use `git log data/southern.csv` and count the number of log entries.'
Ex() >> test_student_typed(r'\s*git\s+log\s+data/southern.csv\s*', \
                           fixed=False, \
                           msg='Use `git log` and the name of a file.') \
     >> test_mc(3, [err, err, 'Correct!', err])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:BulletConsoleExercise key:1be0ce9219
## How do I write a better log message?

Writing a one-line log message with `git commit -m "message"`is good enough for very small changes,
but your collaborators (including your future self) will appreciate more information.
If you run `git commit` *without* `-m "message"`,
Git launches a text editor with a template like this:

```

# Please enter the commit message for your changes. Lines starting
# with '#' will be ignored, and an empty message aborts the commit.
# On branch master
# Your branch is up-to-date with 'origin/master'.
#
# Changes to be committed:
#       modified:   skynet.R
#
```

The lines starting with `#` are comments, and won't be saved.
(They are there to remind you what you are supposed to do and what files you have changed.)
Your message should go at the top, and may be as long and as detailed as you want.

We have configured Git to use the Nano editor for this course.
The final chapter of this course will show you how to set up a different editor.

*** =pre_exercise_code
```{python}
append = '''

Fehrenbach: "Dental Anatomy Coloring Book" (2e), 2013.
'''
with open('dental/report.txt', 'w') as writer:
    writer.write(append)
repl = connect('bash')
repl.run_command('cd dental')
repl.run_command('git add report.txt')
```

*** =type1: ConsoleExercise
*** =key1: 2f3aa2a066

*** =xp1: 10

*** =instructions1

You are in the `dental` repository.
The changes to `report.txt` have already been staged.
Use `git commit` *without* `-m` to commit the changes.

*** =hint1

*** =sample_code1
```{shell}
```

*** =solution1
```{shell}
git commit -m "Adding a reference."
```

*** =sct1
```{python}
Ex() >> test_student_typed(r'\s*git\s+commit\s*',
                           fixed=False,
                           msg='Use `git add` without `-m`.')
```
