---
title       : Basic workflow
description : >-
  This chapter shows you how to use Git to back up your work, and how
  to view the history of what you have done.

--- type:PureMultipleChoiceExercise lang:bash xp:50 key:a3732cc273
## What is version control?

A **version control system** is a tool that manages changes made to the files and directories in a project.
Many version control systems exist;
this lesson focuses on one called Git,
which is used by many of the data science tools covered in our other lessons.
Its strengths are:

* Nothing that is committed to Git is ever lost,
  so you can always go back to see which results were generated by which versions of your programs.

* Git automatically notifies you when your work conflicts with someone else's,
  so it's harder (but not impossible) to accidentally overwrite work.

* Git can synchronize work done by different people on different machines,
  so it scales as your team does.

Version control isn't just for software:
books, papers, data sets, and anything that changes over time or needs to be shared
can and should be stored and shared using something like Git.

<hr>

Which of the following does Git do?

*** =possible_answers
- Keep track of changes to files.
- Notice conflicts between changes made by different people.
- Synchronize files between different computers.
- [All of the above.]

*** =hint

*** =feedbacks
- Yes, but that's not all.
- Yes, but that's not all.
- Yes, but that's not all.
- Correct.

<!-- -------------------------------------------------------------------------------- -->

--- type:PureMultipleChoiceExercise lang:bash xp:50 key:74e878cc1b
## How does Git compare to alternatives?

Git and other version control systems are the only way to track and share work.
Some popular alternatives are:

1. *Email.*
   Lots of people mail files back and forth,
   but it's hard to keep track of which version is the most recent,
   and even harder to be sure that everyone's edits have been merged together.

2. *Putting files in a shared network drive or on Dropbox.*
   This is more convenient than email,
   but still makes people do the work of tracking changes.

3. *Online editing.*
   Tools like Google Docs work well for small groups,
   but it's hard to keep track of changes as the number of contributors and comments grows.

<hr>

What kind of work should you store in version control?

*** =possible_answers
- Software.
- Small datasets.
- Reports and other written work.
- [All of the above.]

*** =hint

*** =feedbacks
- Yes, but that's not all.
- Yes, but that's not all.
- Yes, but that's not all.
- Correct!

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:bc52cf1174
## How can I check the state of a repository?

The previous chapter showed you how to look at a repository's history.
To create that history,
you must save the changes you make to the repository's contents.

The first step is to check the **status** of the repository's files.
To do this,
run the command `git status`,
which displays a list of the files that have been modified
since the last time changes were committed.

<hr>

You are in the `dental` repository.
Use a single Git command to discover which file(s) have been changed since the last commit.
Which file(s) are listed?

*** =instructions
- `data/summer.csv`.
- `report.txt`.
- Neither of the above.
- Both of the above.

*** =hint

*** =pre_exercise_code
```{python}
append = '''

Fehrenbach: "Dental Anatomy Coloring Book" (2e), 2013.
'''
with open('dental/report.txt', 'w') as writer:
    writer.write(append)
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
Ex() >> test_student_typed(r'\s*git\s+status\s*', \
                           fixed=False, \
                           msg='Use `git status`.') \
     >> test_mc(2, ['No, that file has not changed.', \
                    'Correct!', \
                    'No, one file has changed.', \
                    'No, only one file has changed.'])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:BulletConsoleExercise key:55204eee61
## What's the first step in saving changes?

You commit changes to a Git repository in two steps:
you use Git to add one or more files to a **staging area**,
and then you use it to commit everything in the staging area.
Putting files in the staging area is like putting things in a box,
while committing is like putting the box in the mail.

To add a file to the staging area,
using `git add filename`.
You can do this several times in a row as you make changes to the file
without leaving traces in the repository's history.
Once you commit those changes,
though,
they become part of the repository's permanent log.

*** =pre_exercise_code
```{python}
append = '''

Fehrenbach: "Dental Anatomy Coloring Book" (2e), 2013.
'''
with open('dental/report.txt', 'w') as writer:
    writer.write(append)
connect('bash').run_command('cd dental')
```

*** =type1: ConsoleExercise
*** =key1: a34bf017f0

*** =xp1: 10

*** =instructions1

You are in the `dental` repository.
Use `git add` to add the file `report.txt` to the staging area.

*** =hint1

*** =sample_code1
```{shell}
```

*** =solution1
git add report.txt
```{shell}
```

*** =sct1
```{python}
Ex() >> test_student_typed(r'\s*git\s+add\s+report.txt\s*',
                           fixed=False,
                           msg='Use `git add` and a filename.')
```

*** =type2: ConsoleExercise
*** =key2: 961661800c

*** =xp2: 10

*** =instructions2

Use another Git command to check the repository's status.

*** =hint2

*** =sample_code2
```{shell}
```

*** =solution2
```{shell}
git status
```

*** =sct2
```{python}
Ex() >> test_student_typed(r'\s*git\s+status\s*',
                           fixed=False,
                           msg='Use `git status`.')
```

<!-- -------------------------------------------------------------------------------- -->

--- type:BulletConsoleExercise key:f208f45d7d
## How can I tell what's going to be committed?

`git status` shows you which files are in the staging area,
and which files have changes that haven't yet been put there.
In order to compare the file as it currently is
to the changes in the staging area,
you can use `git diff -r HEAD filename`.
(Remember, `HEAD` is a shortcut meaning "the most recent commit".)

*** =pre_exercise_code
```{python}
with open('dental/data/northern.csv', 'w') as writer:
    writer.write('2017-11-01,bicuspid\n')
with open('dental/data/eastern.csv', 'w') as writer:
    writer.write('2017-11-02,molar\n')
repl = connect('bash')
repl.run_command('cd dental')
repl.run_command('git add data/northern.csv')
```

*** =type1: ConsoleExercise
*** =key1: 77c975a5c8

*** =xp1: 10

*** =instructions1

You are in the `dental` repository.
Use one Git command to see which files have staged and unstaged changes.

*** =hint1

*** =sample_code1
```{shell}
```

*** =solution1
```{shell}
git status
```

*** =sct1
```{python}
Ex() >> test_student_typed(r'\s*git\s+status\s*',
                           fixed=False,
                           msg='Use `git status`.')
```

*** =type2: ConsoleExercise
*** =key2: 5a866b3ef0

*** =xp2: 10

*** =instructions2

Use a single Git command to view the changes in the file that has been staged
(and *only* that file),

*** =hint2

*** =sample_code2
```{shell}
```

*** =solution2
```{shell}
git diff -r HEAD data/northern.csv | cat
```

*** =sct2
```{python}
Ex() >> test_student_typed(r'\s*git\s+diff\s+-r\s+HEAD\s+data/northern\.csv\s*',
                           fixed=False,
                           msg='Use `git diff -r HEAD` and a filename.')
```

*** =type3: ConsoleExercise
*** =key3: 538a9c35d4

*** =xp3: 10

*** =instructions3

Use one Git command to add the other changed file to the staging area.

*** =hint3

*** =sample_code3
```{shell}
```

*** =solution3
```{shell}
git add data/eastern.csv
```

*** =sct3
```{python}
Ex() >> test_student_typed(r'\s*git\s+add\s+data/eastern\.csv\s*',
                           fixed=False,
                           msg='Use `git add`.')
```

<!-- -------------------------------------------------------------------------------- -->

--- type:BulletConsoleExercise key:dbf5aa722c
## How do I commit changes?

To save the changes in the staging area,
you use the command `git commit`.
It always saves everything that is in the staging area as one unit:
as you will see later,
when you want to undo changes to a project,
you undo all of a commit or none of it.

When you commit changes,
Git requires you to enter a log message.
As you saw in the previous chapter,
this serves the same purpose as a comment in a program:
it tells the next person to examine the repository why you made a change.

By default,
Git launches a text editor to let you write this message,
but you can also use the flag `-m "some message in quotes"` on the command line
to enter a single-line message like this:

```
git commit -m "Program appears to have become self-aware."
```

*** =pre_exercise_code
```{python}
append = '''

Fehrenbach: "Dental Anatomy Coloring Book" (2e), 2013.
'''
with open('dental/report.txt', 'w') as writer:
    writer.write(append)
repl = connect('bash')
repl.run_command('cd dental')
repl.run_command('git add report.txt')
```

*** =type1: ConsoleExercise
*** =key1: 17e219ea22

*** =xp1: 10

*** =instructions1

You are in the `dental` repository.
Use one Git command to check the status of the repository.

*** =hint1

*** =sample_code1
```{shell}
```

*** =solution1
```{shell}
git status
```

*** =sct1
```{python}
Ex() >> test_student_typed(r'\s*git\s+status\s*',
                           fixed=False,
                           msg='Remember, you want to check the *status* of the repository.')
```

*** =type2: ConsoleExercise
*** =key2: a5ce3bebb8

*** =xp2: 10

*** =instructions2

Commit the changes in the staging area with the message "Adding a reference."

*** =hint2

*** =sample_code2
```{shell}
```

*** =solution2
```{shell}
git commit -m "Adding a reference."
```

*** =sct2
```{python}
Ex() >> test_student_typed(r'\s*git\s+commit\s+-m\s+"Adding\s+a\s+reference."\s*',
                           fixed=False,
                           msg='Use `git commit` with `-m "message"`.')
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:97ae2f8613
## How can I view a repository's entire history?

Every Git command consists of the word `git` followed by a verb that specifies what you want to do.
One of the most common commands is `git log`,
which tells Git to display the **log** of the project's history.
There are two ways to run it:

1. If you are in a project's root directory or one of its sub-directories,
   you can just run `git log`.
2. Otherwise, you can run `git -C /path/to/repository log`.
   The flag `-C /some/path` works with every Git command,
   which is why it comes between the word `git` and the specific command.

Log entries are shown most recent first,
and look like this:

```
commit 0430705487381195993bac9c21512ccfb511056d
Author: Rep Loop <repl@datacamp.com>
Date:   Wed Sep 20 13:42:26 2017 +0000

    Added year to report title.
```

Ignoring the `commit` line for a moment,
the other lines show who made the change,
when,
and what comment (called a **log message**) they wrote for the change.

<hr>

You are in the directory `dental`,
which is a Git repository.
Use a single Git command (*without* the `-C` option)
to view the repository's history.
What is the message on the very first entry in the log
(which is displayed last)?

*** =instructions
- "Added summary report file."
- "Adding seasonal CSV data files."
- "Fixed bug and generated results."
- "Adding reminder to cite funding sources."

*** =hint

Use spacebar to page down through the log and 'q' to quit.

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
err = 'Incorrect: please re-run the command and use `git log`.'
Ex() >> test_mc(1, ['Correct', err, err, err])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:77aa525e25
## How can I view a specific file's history?

A project's entire log can be overwhelming,
so it's often useful to inspect only the changes to particular files or directories.
You can do this using `git log path`,
where `path` is the path to a specific file or directory.
The log for a file shows changes made to that file;
the log for a directory shows when files were added or deleted in that directory,
rather than when the contents of the directory's files were changed.

<hr>

You are in the `dental` repository.
Use `git log` to display only the changes made to `data/southern.csv`.
How many have there been?

*** =instructions
- 0.
- 1.
- 2.
- 3.

*** =hint

Use `git log` and the path to the file you are interested in.

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
err = 'Incorrect: please use `git log data/southern.csv` and count the number of log entries.'
Ex() >> test_student_typed(r'\s*git\s+log\s+data/southern.csv\s*', \
                           fixed=False, \
                           msg='Use `git log` and the name of a file.') \
     >> test_mc(3, [err, err, 'Correct!', err])
```

<!-- -------------------------------------------------------------------------------- -->

--- type:BulletConsoleExercise key:1be0ce9219
## How do I write a better log message?

Writing a one-line log message with `git commit -m "message"`is good enough for very small changes,
but your collaborators (including your future self) will appreciate more information.
If you run `git commit` *without* `-m "message"`,
Git launches a text editor with a template like this:

```

# Please enter the commit message for your changes. Lines starting
# with '#' will be ignored, and an empty message aborts the commit.
# On branch master
# Your branch is up-to-date with 'origin/master'.
#
# Changes to be committed:
#       modified:   skynet.R
#
```

The lines starting with `#` are comments, and won't be saved.
(They are there to remind you what you are supposed to do and what files you have changed.)
Your message should go at the top, and may be as long and as detailed as you want.

> To keep things simple,
> we have configured Git to use the Nano editor for this course.
> The final chapter of this course will show you how to set up a different editor.

*** =pre_exercise_code
```{python}
append = '''

Fehrenbach: "Dental Anatomy Coloring Book" (2e), 2013.
'''
with open('dental/report.txt', 'w') as writer:
    writer.write(append)
repl = connect('bash')
repl.run_command('cd dental')
repl.run_command('git add report.txt')
```

*** =type1: ConsoleExercise
*** =key1: 2f3aa2a066

*** =xp1: 10

*** =instructions1

You are in the `dental` repository.
The changes to `report.txt` have already been staged.
Use `git commit` *without* `-m` to commit the changes.

*** =hint1

*** =sample_code1
```{shell}
```

*** =solution1
```{shell}
git commit -m "Adding a reference."
```

*** =sct1
```{python}
Ex() >> test_student_typed(r'\s*git\s+commit\s*',
                           fixed=False,
                           msg='Use `git add` without `-m`.')
```
