--- type:MultipleChoiceExercise lang:shell xp:50 skills:1 key:d03c5fa3e6
## How can I view a specific log entry?

Every commit to a repository has a unique hexadecimal ID called a **hash**
(since it is generated by running the changes through a pseudo-random number generator
called a **hash function**).
To view the details of a specific commit,
you can use the command `git show` with the first few characters of the commit's hash.
(5 or 6 characters is usually enough to tell Git which hash you mean.)
For example, the command `git show 043070` produces this:

    commit 0430705487381195993bac9c21512ccfb511056d
    Author: Rep Loop <repl@datacamp.com>
    Date:   Wed Sep 20 13:42:26 2017 +0000
    
        Added year to report title.
    
    diff --git a/report.txt b/report.txt
    index e713b17..4c0742a 100644
    --- a/report.txt
    +++ b/report.txt
    @@ -1,4 +1,4 @@
    -# Seasonal Dental Surgeries 2017-18
    +# Seasonal Dental Surgeries (2017) 2017-18
     
     TODO: write executive summary.

The first part is the log entry you have seen before.
The second part shows what changes the commit made:
sections that the change removed are prefixed with `-`,
while sections that it added are prefixed with `+`.

<hr>

How many files were changed by the commit whose hash starts with `aa5991c6`?
(Remember, you only need to type the first few characters of the hash.)

*** =instructions
- None.
- 1.
- 2.
- 4.

*** =hint

Use `git show` and the first 6 characters of the hash
and then look at the names of the files that were changed.

*** =pre_exercise_code
```{shell}
repl = connect('bash')
repl.run_command('cd dental')
```

*** =sct
```{python}
e_more = 'No, there have been more changes than that.'
correct = 'Correct!'
e_fewer = 'No, there have been fewer changes than that.'
Ex() >> test_student_typed(r'\s*git\s+show\s+4d3e[0-9a-f]*\s*', \
                           fixed=False, \
                           msg='Use `git show` and the first few digits of the commit hash.') \
     >> test_mc(3, [e_more, e_more, correct, e_fewer])
```

